# Lab: Hypothesis Testing


## More Sleep Always! 

SLEEP! It's a beautiful thing that most of us don't get enough of!  There has been a lot of research done looking the amount of sleep different people get on average.  The following is an example of one of them. 

In 2013, [Gallup completed their latest survey on sleep](https://news.gallup.com/poll/166553/less-recommended-amount-sleep.aspx) that shows that the average amount of sleep Americans get is 6.8 hours ($\mu = 6.8$). This seems a bit low, right? How do you think our class average in Data Science DISCOVERY compares to this? It might be less than 6.8 or it may be more than that. Let's find out if the ideal hours of sleep STAT 107 students would get is significantly different from the average amount of sleep that all Americans get reported by Gallup. For the sake of this lab, assume our survey data is a random sample from the population (Americans).

Run a **hypothesis test** to find out if the data we have is significantly different from the population/Gallup poll. 

You need to look at the Day 1 Dataset (from the first week of class) which gives you the ideal amount of sleep Stat 107 students get each week. Here's the link (make sure to login with @illinois): https://docs.google.com/spreadsheets/d/1LB4E89GoVi5yezEfHhopbI5BMxodSG7VvoznkaAkMAQ/edit#gid=0

As always, start with uploading dataset to your notebook, and import the relevant libraries to your notebook.


```python
# Import the libraries
import pandas as pd
import random
from scipy.stats import norm

## Import your data set
df = pd.read_csv("hello.csv")
df
```

Output:
```
Name                      Major School Year     Type  \
0         Wade           Computer Science       Other  Android   
1        Karle                 Statistics       Other   iPhone   
2         noor  global studies/statistics   Sophomore   iPhone   
3     Jonathan          Political Science      Senior   iPhone   
4        Clare                 Statistics   Sophomore   iPhone   
..         ...                        ...         ...      ...   
161  Stephanie          Political Science      Junior   iPhone   
162      Cindy     biomedical engineering      Senior  Android   
163      Talal       Information sciences   Sophomore   iPhone   
164       John     Statistics or Business    Freshman   iPhone   
165     Anisha                        Mcb      Junior   iPhone   

                    Computer  Straw Holes  Has Dogs  Has Cats  Has Fish  \
0     Windows-based computer          2.0         0         1         0   
1    Mac OS X-based computer          1.0         0         0      
```

## Puzzle 1: Sleep Among STAT 107 Students

### Part 1: Null and Alternative Hypothesis
In class, we learned that the first step when doing any hypothesis test is to state the **Null Hypothesis** and **Alternative Hypothesis**. 

**Null Hypothesis** := The null hypothesis says that any difference between our sample and the population is small and due to chance.

**Alternative Hypothesis** := The alternative hypothesis says that the differences we observe are not just due to chance. 

_Warning!!!_ Remember that when you are stating your hypotheses, they always include the **population parameters** such as $\mu$ not **sample parameters** $\bar{X}$. 

Using the Day 1 dataset and the Gallup poll, state your 2 hypotheses:

### Null Hypothesis

$H_0$ := *EDIT THIS to define the null hypothesis*


### Alternative Hypothesis

$H_a$ := *EDIT THIS to define the alternative hypothesis*


```python
# The null hypothesis is that there is no difference between the 2 different groups and they both sleep the same

# The alternative hypothesis is 6.8 hours of sleep is not what college students get on average
```

### Part 2: P-Values:
In lecture, Karle and Wade discussed what a **p-value** is.  We can use it to decide which hypothesis is supported by our data.

Here, it will tell us the probability that we will observe the difference we observed between our sample data and the Gallup data, given that the true average hours of sleep is really 6.8 hours, as reported by Gallup.

Calculate the **p-value** in the following cell:

```python
avg= df['Hours of Sleep'].mean()
avg
```

Output:
```
7.051204819277109
```

```python
sd= df['Hours of Sleep'].std()
sd
```

Output:
```
0.8818532793382718
```

```python
se= sd/(len(df['Hours of Sleep']))**(1/2)
se
```

Output:
```
0.06844508863542532
```

```python
# find first your z value
z_test = (6.8-avg)/(se)
z_test
```

Output:
```
-3.6701657384821034
```

<!-- grade_id: cell-dc78f2d3d7118a28, points: 0, grade: false, solution: true -->
```python
# find your p value
p_value = 2* norm.cdf(z_test)
p_value

```

Output:
```
0.00024239325252541324
```

<!-- grade_id: cell-1a8a95387b8f9d9f, points: 1, grade: true, solution: false -->
```python
# == TEST CASES for Puzzle 1 ==
# - This read-only cell contains test cases for your previous cell(s).
# - If this cell runs without any errors in the output, you PASSED all test cases!
# - If this cell results errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
assert((z_test- 3.590265738)< 0.08), "You z statistics does not seem correct. Make sure you pick the same formula"
assert((p_value - 0.000001 < 0.0013) & (0.0002 < (p_value - 0.000001))), "You p-value statistics does not seem correct."



## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()
```

Output:
```
🎉 All tests passed! 🎉
```

### Part 3: Connect the p value result with your original hypotheses :

We know what the **p-value** is now! But as a data scientist, it's important to explain your results in a way that anyone (even someone who has no background in statistics) can understand. If we just said "Here is the p-value", they may be confused. We need to be more specific about our conclusion. 

In the following cell, delete the statements that you do not believe are true and keep the one you is correct.


<!-- grade_id: cell-0f84432fe81139e8, points: 0, grade: false, solution: true -->
```python
# Remove the of the following lines to leave only the correct answer:

conclusion = "There is enough evidence to reject the null hypothesis. Thus, DO reject the Null"


```

<!-- grade_id: cell-119c28992fed3d21, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1 ==
# - This read-only cell contains test cases for your previous cell(s).
# - If this cell runs without any errors in the output, you PASSED all test cases!
# - If this cell results errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
assert(len(conclusion) == 80), "Your conclusion does not seem correct. Think about more!"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

Output:
```
🎉 All tests passed! 🎉
```

## Puzzle 2: Who loves sleep more?

We are curious if there is a significant difference in ideal hours of sleep between introverts and ambiverts in college. We are specifically interested in whether introverts sleep **less than** ambiverts.

We can investiage this using our Day 1 Dataset as well!  Here, assume that the Day 1 Dataset is a random sample taken from all college students.

Do a hypothesis test to satisfy our curiosity! As awlays, first you need to state your hypotheses.

### Part 1: Null and Alternative Hypothesis

### Null Hypothesis

The null is that we believe that there is no difference between the amount of sleep that introverts and ambiverts get

### Alternative Hypothesis
The alternative is that we believe that introverts sleep less than ambiverts

### Part 2: P-value:

Calculate your test statistic and find the p-value. Remember, we are comparing two samples now- think about how this is different from Puzzle 1 and what type of z-test to do!

Find the **p-value** in the following cell: 

```python
df
```

Output:
```
Name                      Major School Year     Type  \
0         Wade           Computer Science       Other  Android   
1        Karle                 Statistics       Other   iPhone   
2         noor  global studies/statistics   Sophomore   iPhone   
3     Jonathan          Political Science      Senior   iPhone   
4        Clare                 Statistics   Sophomore   iPhone   
..         ...                        ...         ...      ...   
161  Stephanie          Political Science      Junior   iPhone   
162      Cindy     biomedical engineering      Senior  Android   
163      Talal       Information sciences   Sophomore   iPhone   
164       John     Statistics or Business    Freshman   iPhone   
165     Anisha                        Mcb      Junior   iPhone   

                    Computer  Straw Holes  Has Dogs  Has Cats  Has Fish  \
0     Windows-based computer          2.0         0         1         0   
1    Mac OS X-based computer          1.0         0         0      
```

```python
df_introvert= df[df['Personality']=='Introvert']
df_introvert
```

Output:
```
Name                                              Major  \
0              Wade                                   Computer Science   
2              noor                          global studies/statistics   
3          Jonathan                                  Political Science   
5             Dhruv                                   Computer science   
6              Alex                                            Physics   
10            Taige                                          Math + CS   
14             Gail                                        Art History   
22             Katy                                         Statistics   
26             Eric                                       Not sure yet   
27             Mike                                         Statistics   
28            Chloe                             Structural Engineering   
29           Eugene                      Statistics + Computer Science   
30        Folashade                                  
```

```python
mean_Introvert = df_introvert['Hours of Sleep'].mean()
print(mean_Introvert)

sd_Introvert= df_introvert['Hours of Sleep'].std()
print(sd_Introvert)
```

Output:
```
6.875
0.9095771041943482
```

<!-- grade_id: cell-d45b5ab2db34da59, points: 0, grade: false, solution: true -->
```python
## Make sure your mean_Introvert; sd_Ambivert; pooled_sd; p_value names as below for the test cases
df_ambivert= df[df['Personality']=='Ambivert']

mean_ambivert= df_ambivert['Hours of Sleep'].mean()
print(mean_ambivert)

sd_Ambivert = df_ambivert['Hours of Sleep'].std()
print(sd_Ambivert)

```

Output:
```
7.160919540229885
0.8676828701884063
```

```python
se_introvert= sd_Introvert/ (len(df_introvert))**(1/2)
se_introvert
```

Output:
```
0.1174258992212093
```

```python
se_ambivert= sd_Ambivert/ (len(df_ambivert))**(1/2)
se_ambivert
```

Output:
```
0.09302536813930255
```

```python
se_difference = ((se_ambivert)**(2)+(se_introvert)**(2))**(1/2)
se_difference
```

Output:
```
0.14980841406730921
```

```python

```

```python
p_value = (mean_Introvert - mean_ambivert)/se_difference
p_value
```

Output:
```
-1.9085679666925812
```

<!-- grade_id: cell-7bd5b9e9b0bf72b8, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1 ==
# - This read-only cell contains test cases for your previous cell(s).
# - If this cell runs without any errors in the output, you PASSED all test cases!
# - If this cell results errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
assert(mean_Introvert - 0.00001 < 6.9 ), "Your mean of Introvert does not look like correct. Make sure you correctly subset your data"
assert(sd_Ambivert- 0.00001 <0.90 ), "Your standard error of Ambivert does not look like correct. Make sure you correctly subset your data"
assert(se_difference - 0.0001 < 0.2 ), "Your standard error for the difference does not look like correct."
assert(p_value - 0.0001 < 0.03 ), "Your p-value does not look like correct. Make sure you are doing one sided test."


## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

Output:
```
🎉 All tests passed! 🎉
```

### Part 3: Make a conclusion based on your p-value.

We know what the **p-value** is now. So in the following cell delete the statements that you do not believe are true and keep the one you think is true.  Use 1% as your cutoff instead of 5%.

<!-- grade_id: cell-5f2236ee09b8ae71, points: 0, grade: false, solution: true -->
```python
# Comennt the of the following lines to leave only the correct answer:
#conclusion = "There is enough evidence to reject the null hypothesis. Thus, DO reject the Null"
#conclusion = "We cannot infer anything!"
conclusion = "There is not enough evidence to reject the null hypothesis. Thus, DO NOT reject the Null"


```

<!-- grade_id: cell-5ba76e5cb1428aa4, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1 ==
# - This read-only cell contains test cases for your previous cell(s).
# - If this cell runs without any errors in the output, you PASSED all test cases!
# - If this cell results errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
assert(len(conclusion) == 88), "Your conclusion does not seem correct. Think about more!"

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()

```

Output:
```
🎉 All tests passed! 🎉
```

## Puzzle 3: Halloween!

![pumpkins.jpeg](pumpkins.jpeg)

Karle heard that 20% of college students dress up in the same costume for two consecutive Halloweens. However Karle thinks that this percent is higher for UIUC students!  She thinks it is **40%** and wants to test this hypothesis. She decides to take 1000 differents samples of 60 students this Hallowen.

We can simulate this in Python. 

Can you help her simulate this and test her idea by finding an approximate p-value. As always, start with the null and alternative hypotheses!

Here are the steps that can help you: 

1. Consider a very large sample that can simulate an example population -- use `random.randint(...)`.

2. Randomly sample 60 students 1000 times. For each of these 1000 samples, find the percentage that plan to wear the same costume from last halloween. 

3. You have 1000 different samples and sample percentages for all of them. Now think about how you can the p-value.  Hint: The p-value is the probability of getting 40% or more people wearing the same costume given the true percentage is 20%.


<!-- grade_id: cell-d3610d77d94bdcbf, points: 0, grade: false, solution: true -->
```python
## Create the DataFrame of 100 samples:
def population():
    total= 0
    for i in range(60):
        if random.randint(1,60) <= 12:
            total= total + 1
    return total/60


data= []
for i in range(1000):
    data.append({'panel':population()})

df= pd.DataFrame(data)

df
```

Output:
```
panel
0    0.200000
1    0.166667
2    0.183333
3    0.166667
4    0.183333
..        ...
995  0.250000
996  0.166667
997  0.216667
998  0.266667
999  0.216667

[1000 rows x 1 columns]
```

```python
observed= df.mean()
```

```python
## make sure your answer is the p-value
p_value = (observed- .40)/(((.2*.8)**(1/2))/(len(df)**(1/2)))**(1/2)
p_value =p_value[0]

```

<!-- grade_id: cell-c97db47b822b8ecc, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1 ==
# - This read-only cell contains test cases for your previous cell(s).
# - If this cell runs without any errors in the output, you PASSED all test cases!
# - If this cell results errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
assert(len(df) >= 1000), "You must heave at least 1000 samples"
assert(p_value <= 0.02), "Your p value does not seem correct!"


## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()
```

Output:
```
🎉 All tests passed! 🎉
```

## Puzzle 4: GPA Visualization

Using Wade's [Grade disparity between sections at UIUC](http://waf.cs.illinois.edu/discovery/grade_disparity_between_sections_at_uiuc/) visualization, find the average grade Wade gave to CS 225 students.

You suspect Wade might not be completely honest with his grades, so you randomly survey 107 students who have taken CS 225.  The final grades the students gave you is out of 1000 points.  You have recorded the 107 student grades in `cs225-final-grades.csv`.  (*These are really 107 randomly chosen grades from CS 225 from Fall 2018.*)

You find that the translation between points and grades is as follows:

| Points | Grade | GPA |
| ------ | ----- | --- |
| 1070 - 1100 | A+ | 4.0 |
|  930 - 1070 | A  | 4.0 |
|  900 - 930 | A-  | 3.67 |
|  870 - 900 | B+  | 3.33 |
|  830 - 870 | B   | 3.0 |
|  800 - 830 | B-  | 2.67 |
|  770 - 800 | C+  | 2.33 |
|  730 - 770 | C   | 2.0 |
|  700 - 730 | C-  | 1.67 |
|  670 - 700 | D+  | 1.33 |
|  630 - 670 | D   | 1.0 |
|  600 - 630 | D-  | 0.67 |
|    0 - 600 | F  | 0 |

**Hint**: You will need to process `cs225-final-grades.csv` and understand and gather data from the GPA visualization before you will be able to run the test.  Use `pandas` to help you process this data! :)

<!-- grade_id: cell-053f4c444b109148, points: 0, grade: false, solution: true -->
```python
import pandas as pd
## First read the data cs225-final-grades.csv
sample_cs225 = pd.read_csv("cs225-final-grades.csv")
converted_grades= sample_cs225
converted_grades['GPA']= 0.0

## Procecess sample_cs225 to convert to GPA format. Make sure conv_sample_cs225 is a DataFrame
for i in range(107):
    if sample_cs225['CS 225 Final Grade'][i]>=1070:
        converted_grades['GPA'][i]=4.0
    elif sample_cs225['CS 225 Final Grade'][i]>=930:
        converted_grades['GPA'][i]=4.0
    elif sample_cs225['CS 225 Final Grade'][i]>=900:
        converted_grades['GPA'][i]=3.67
    elif sample_cs225['CS 225 Final Grade'][i]>=870:
        converted_grades['GPA'][i]=3.33
    elif sample_cs225['CS 225 Final Grade'][i]>=830:
        converted_grades['GPA'][i]=3.0
    elif sample_cs225['CS 225 Final Grade'][i]>=800:
        converted_grades['GPA'][i]=2.67
    elif sample_cs225['CS 225 Final Grade'][i]>=770:
        converted_grades['GPA'][i]=2.33
    elif sample_cs225['CS 225 Final Grade'][i]>=730:
        converted_grades['GPA'][i]=2.0
    elif sample_cs225['CS 225 Final Grade'][i]>=700:
        converted_grades['GPA'][i]=1.67
    elif sample_cs225['CS 225 Final Grade'][i]>=670:
        converted_grades['GPA'][i]=1.33
    elif sample_cs225['CS 225 Final Grade'][i]>=630:
        converted_grades['GPA'][i]=1.0
    elif sample_cs225['CS 225 Final Grade'][i]>=600:
        converted_grades['GPA'][i]=.67
    else: converted_grades['GPA'][i]=.00





```

```python
converted_grades
```

Output:
```
CS 225 Final Grade   GPA
0            658.204335  1.00
1           1068.696078  4.00
2           1040.655631  4.00
3           1047.147059  4.00
4            763.440841  2.00
..                  ...   ...
102          954.931372  4.00
103          899.918937  3.33
104          765.717112  2.00
105          328.562675  0.00
106         1048.897892  4.00

[107 rows x 2 columns]
```

<!-- grade_id: cell-0544cc08b9b64325, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1 ==
# - This read-only cell contains test cases for your previous cell(s).
# - If this cell runs without any errors in the output, you PASSED all test cases!
# - If this cell results errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
assert(len(converted_grades[converted_grades["GPA"] == 3.67]) == 7), "Your convertion from the total points to GPA does not seem correct."
assert(len(converted_grades[converted_grades["GPA"] == 0]) == 7), "Your convertion from the total points to GPA does not seem correct."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()
```

Output:
```
🎉 All tests passed! 🎉
```

### Part 1: Hypotheses

Just like before, start with stating your Null and Alternative Hypotheses!

### Null Hypothesis

$H_0$ := There is no difference between the grades that Wade gives and peoples actual scores

### Alternative Hypothesis

$H_a$ := There is a difference between the grades and the actual scores

### Part 2: P-value
Calculate the test statistic and find the p-value in the following cell:

<!-- grade_id: cell-26dab0ff60b63979, points: 0, grade: false, solution: true -->
```python
from scipy.stats import norm
# find first your z value
z_test = 1.50

# find your p value
p_value = .13
#norm.cdf(z_test)


```

<!-- grade_id: cell-08a0447915dbc44c, points: 1, grade: true, solution: false -->
```python
# == TEST CASES for Puzzle 4 ==
# - This read-only cell contains test cases for your previous cell(s).
# - If this cell runs without any errors in the output, you PASSED all test cases!
# - If this cell results errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
assert(round(z_test, 2) == 1.50), "You z statistics does not seem correct. Make sure you pick the same formula"
assert(round(p_value, 2) == 0.13), "You p-value statistics does not seem correct."



## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()
```

Output:
```
🎉 All tests passed! 🎉
```

## Submit Your Work!

You're almost done -- congratulations!

You need to do a few more things:

1. Save your work.  To do this, create a **notebook checkpoint** by using the menu within the notebook to go **File -> Save and Checkpoint**

2. Choose `File` and then `Close and Halt` from this notebook.

3. Choose `Quit` on the main notebook webpage.

4. Return to your command line and follow the directions on the Data Science Discovery page on how to use git to turn this notebook into the course!